<figure>
  <IMG SRC="https://upload.wikimedia.org/wikipedia/commons/thumb/d/d5/Fachhochschule_Südwestfalen_20xx_logo.svg/320px-Fachhochschule_Südwestfalen_20xx_logo.svg.png" WIDTH=250 ALIGN="right">
</figure>

# Machine Learning
### Sommersemester 2022
Prof. Dr. Heiner Giefers

# Erste Schritte in Python

## Variablen

Python ist fundamental objektorientiert. Das heißt nicht nur, dass Sie in Python objektorientiert programmieren können. Es gilt auch der Leitsatz: *Alles in Python ist ein Objekt*. Also selbst grundlegende Datentypen wie `int`, `float` und `str`, sowie Funktionen sind Objekte. Daher sind Variablen in Python immer Referenzen auf Objekte.

Wie viele andere Skriptprachen auch, ist Python dynamisch typisiert. Das bedeutet, dass Sie keine Typangaben bei der Definition einer Variablen angeben müssen. Python leitet automatisch den passenden Typ ab, bzw. es wählt den "am besten passenden" Typ aus.

Sie können daher Variablen in Python wie folgt definieren:

```python
a = 42
b = 1.23
ab = 'Hallo'
```

Für Variablennamen gelten die Gleichen Regeln wie für die Bezeichner in C/C++. Variablen müssen mit einem Buchstaben oder Unterstrich beginnen und können sich ab dem 2. Zeichen aus einer beliebigen Folge von Buchstaben, Ziffern und Unterstrichen zusammensetzen.

Es gibt allerdings einige Konventionen für die Wahl von Bezeichnern. So gelten Variablen, die mit 2 Unterstrichen beginnen als *privat*, Namen mit 2 Unterstrichen am Anfang und am Ende sind für spezielle Attribute und Methoden reserviert ("*magic methods*").

Die skalaren, also elementare oder nicht-zusammengesetzten Datentypen in Python sind:
- `int` Ganze Zahlen
- `float` Fließkommazahlen mit 64-bit Präzision
- `complex` Komplexe Zahlen
- `bool` Bool'scher Datentyp mit den Werten `True` und `False`
- `NoneType` signalisiert das Nicht-Vorhandensein einer Referenz, ähnlich zu `NULL` oder `NIL` in anderen Sprachen

Die Typen `str` für Zeichenketten sowie `bytes` für Folgen von 8-bit (vorzeichenlosen) Werten (zum Verarbeiten von Binärdaten) gehören zu den *Sequentiallen Datentypen*.

## Operationen

Für die meisten Datentypen existieren die bekannten Operationen (`+`, `-`, `*`, `/`) mit der üblichen Bedeutung. Daneben gibt es noch den Operator `//`, für die ganzzahlige Division, den Modulo-Operator `%` und den Potenz-Operator `**`.

```python
a = 2 + 1.23
b = 22.2//3
c = "Hallo " + "Welt"
d = 2**8
print(a, b, c, d)
```

Die `print` Funktion, wie oben verwendet, benötigt man ziemlich häufig.
Ruft man sie mit einer (beliebigen) Folge von Parametern aus, so wird für jeden Variable, entsprechend ihres Typs, eine passende *print* Methode aufgerufen. In Python heißt die Methode `__str__()`, sie entspricht in etwa der `toString()`-Methode aus Java.

Um eine formatierte Ausgabe zu erhalten, kann man einen Format-String mit Platzhaltern angeben, ähnlich wie bei der `printf`-Funktion aus C. Über den Modulo-Operator können dann die Variablen angegeben werden, die an den Platzhaltern eingesetzt werden sollen.
Für unser Beispiel oben sieht das dann z.B. so aus:

```python
print("a = %f, b = %d, c = %s, d = %s" % (a, b, c, d))
```

Python ist stark typisiert. D.h., dass Variablen immer einen eindeutigen Typ haben und an keiner Stelle eine implizite Typumwandlung stattfinden kann. Jede Änderung des Typs erfordert eine explizite Typkonvertierung. Ein Ausdruck wie `"Hallo"+2` kann nicht ausgewertet werden, da die `+`-Operation für einen String und einen Integer nicht definiert ist.
In diesem Fall kann man eine Typenwandlung, z.B. von `int` nach `str` vornehmen:

```python
"Hallo" + str(2)
```

## Sequentielle Datentypen

Unter sequenziellen Datentypen wird eine Klasse von Datentypen zusammengefasst, die Folgen von **gleichartigen oder verschiedenen Elementen** verwalten.
In Listen und Tupel können beliebige Folgen von Daten abgelegt sein. Die gespeicherten Elemente haben eine definierte Reihenfolge und man kann über eindeutige Indizes auf sie zugreifen. Listen sind veränderbar, d.h., man kann einzelne Elemente  ändern, löschen oder hinzufügen. Tupel sind nicht veränderbar. Das bedeutet, bei jeder Änderung wird ein komplett neues Objekt mit den geänderten Elmenten angelegt.

```python
a = [3.23, 7.0, "Hallo Welt", 256]
b = (3.23, 7.0, "Hallo Welt", 256)
print("Liste a = %s\nTupel b =%s" % (a,b) )
print("Das dritte Element von b ist " + b[2])
print("Gleiche Referenz? %s. Gleicher Inhalt? %s" % (a == b, set(a)==set(b)))
```

Das obige Beispiel bringt uns direkt zum nächsten Datentyp, den Mengen (oder engl. *sets*).
Wie bei den Mengen aus der Mathematik kann ein set in Python jedes Objekt nur einmal enthalten.
Wenn wir also aus der Liste [4,4,4,4,3,3,3,2,2,1] eine Menge machen, hat diese folgende Elemente:

```python
set([4,4,4,4,3,3,3,2,2,1])
```

Die elemente tauchen nun nicht nur einmalig auf, sondern sie sind auch umsortiert worden.
Man darf sich hier nicht täuschen lassen, die Elemente einer Menge sind immer unsortiert.
D.h., man kann keine spezielle Sortierung erwarten, auch wenn die Ausgabe in manchen Fällen danach aussieht.

Ein weitere sequentieller Datentyp sind Dictionaries (die deutsche Übersetzung *Wörterbücher* passt hier nicht so gut).
Diectionaries sind eine Menge von *Schlüssel-Wert-Paaren*.
Das bedeutet, dass jeder Wert im Dictionary unter einem frei wählbaren Schlüssel abgelegt ist, und auch über diesen Schlüssel zugegriffen werden kann.

```python
haupstaedte = {"DE" : "Berlin", "FR" : "Paris", "US" : "Washington", "CH" : "Zurich"}
print(haupstaedte["FR"])
haupstaedte["US"] = "Washington, D.C."
print(haupstaedte["US"])
```

## Funktionen 

Funktionen in Python werden über das Schlüsselwort `def` definiert. Die Syntax einer Funktions-Definition sieht folgendermaßen aus:

```python
def myfunc(arg1, arg2,... argN):  
  '''Dokumentation'''  

  #Programmcode  

  return <Rückgabewert>  
```

Hier wird die Funktion "myfunc" definiert, welche mit den Parametern "arg1,arg2,....argN" aufgerufen werden kann.

Wir sehen hier auch ein weiteres Konzept von Python, das wir bisher noch nicht angesprochen haben. Die Strukturierung von Code in Blöcke erfolgt über **Einrückungen**.
Für eine Funktion bedeutet das, dass der Code des Funktionskörpers um eine Stufe gegenüber der Funktionsdefinition eingerückt sein muss. Wenn der Funktionskörpers weitere Kontrollstrukturen enthält, z.B. Schleifen oder Bedingungen, sind weitere Einrückungen nötig. Betrachten Sie folgendes Beispiel:

```python
def gib_was_aus():
    print("Eins")
    print("Zwei")
print("Drei")
gib_was_aus()
```

Hier wird zuerst eine Funktion `gib_was_aus` definiert. Die Anweisung `print("Drei")` ist nicht mehr eingerückt, gehört daher nicht mehr zur Funktion.

Funktionen können fast überall definiert sein, also z.B. auch innerhalb von anderen Funktionen.
Rückgaben erfolgen, wie auch in anderen Programmiersprachen üblich mit den Schlüsselwort `return`.
Falls mehrere Elemente zurückgegeben werden sollen, können diese z.B. in ein Tupel gepackt werden:

```python
def inc(a, b, c):
    return a+1, b+1, c+"B"
a=b=1
c="A"
a,b,c = inc(a,b,c)
print(a,b,c)
```

## Verzweigungen

Wir haben bisher noch keine Kontrollstrukturen, also Verzweigungen oder Schleifen angesprochen.
Eine Bedingung oder Verzeigung funktioniert in Python (wie üblich) über ein `if`-`else`-Konstrukt.
Auch hier werden zur Strukturierung der Blöcke Einrückungen benutzt.

```python
a=2
if a==0:
    print("a ist Null")
else:
    print("a ist nicht Null")

```

Um tiefe Verschachtelungen zu vermeiden, gibt es noch ein `elif`-Anweisung:

```python
a=2
if a<0:
    print("a ist negativ")
elif a>0:
    print("a ist positiv")
else:
    print("a ist Null")
```

## Schleifen

In Python gibt es die Schleifentypen, `while` und `for`, wobei letztere eine etwas ungewöhnliche Syntax hat.
Die `while`-Schleife hingegen wird wie in vielen bekannten Programmiersprachen benutzt:

```python
i = 5
while i>0:
    print(i)
    i -= 2
```

Anders als z.B. in C/C++ oder Java läuft eine `for`-Schleife in Python nicht über eine Zählvariable, sondern über *die Elemente eines iterierbaren Datentyps*.
Einige Beispiele für iterierbaren Datentypen haben wir schon als sequentielle Datentypen kennen gelernt.
Wir können z.B. mit einer `for`-Schleife alle Elemente eines Dictionaries besuchen:

```python
haupstaedte = {"DE" : "Berlin", "FR" : "Paris", "US" : "Washington", "CH" : "Zurich"}
for s in haupstaedte:
    print(haupstaedte[s])
```

Wir sehen, dass die Laufvariable hier alle Schlüssel des Dictionaries annimmt. Bei einer Liste wird über alle Werte iteriert:

```python
a = [3.23, 7.0, "Hallo Welt", 256]
for s in a:
    print(s)
```

Neben den sequentiellen Datentypen liefern noch sogenannte **Generatoren** Folgen von Werten die iterierbar sind. Der bekannteste iterator ist `range()`.
`range` kann mehrere Argumente haben. Ist nur ein Argument `E` angegeben, so läuft der iterator von 0 bis `E-1`.
`range(S, E)` läuft von S bis `E-1`, und `range(S, E, K)` läuft von S bis `E-1` mit der Schrittweite `K`

```python
print("Ein Parameter:", end=" ")
for s in range(5): print(s, end=" ")

print("\nZwei Parameter:", end=" ")
for s in range(2,5): print(s, end=" ")

print("\nDrei Parameter:", end=" ")
for s in range(0,5,2): print(s, end=" ")
```

Das zusätzliche Argument `end=" "` in den `print`-Anweisungen oben verhindert übrigens einen Zeilenumbruch.
Ohne diesen Parameter würden alle Werte in einer Spalter untereinander ausgegeben.

Damit endet unser erster *Crash Kurs* zum Thema Python. Sie haben nun die wichtigsten Elemente der Python-Syntax gesehen.
Natürlich zeigen die Beispiele aber nur einen kleinen Auschnitt, die Sprache ist noch deutlich umfangreicher und viele Konzepte, wie z.B. Klassen und Module, haben wir noch nicht einmal angesprochen.

Am besten Sie probieren Python einfach mal aus, indem Sie bestehende Beispiele übernehmen und verändern.
Die Python Notebooks sind eine ideale Umgebung dafür.
Sie können in den Code-Zellen Programmcode einfach ausprobieren.
In den Markdown-Zellen können Sie sich Notizen machen, um Ihren Code zu dokumentieren oder ihre Schritte zu beschreiben.

## Aufgaben

Auf den folgenden Notebooks wird es ggf. auch Aufgaben zur Eigneständigen Bearbeitung geben. Ihre Lösungen können Sie über das JupyterNotebook einreichen.

Zum Testen der Abgaben kommt hier eine Mini Aufgabe:
Implementieren Sie eine Python Funktion, die den String *Hello World* zurückliefert!

<!-- grade_id: cell-9386ee03d73035d5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6f798c39f9adb356, points: 1, grade: true, solution: false -->
```python
assert hello()=="Hello World"
```